---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise that David came up with.

`@instructions`
print the result of 8 + 10 with the print statement.

`@hint`
use the print statement!!!!!!!

`@pre_exercise_code`
```{python}
import numpy
```

`@sample_code`
```{python}
print()
```

`@solution`
```{python}
print(8+10)
```

`@sct`
```{python}
Ex().has_output(18, pattern=True)
```
